---
title: "Week 3 Lab"
output:
  learnr::tutorial:
    progressive: true
    allow_skip: true
    fig_caption: false
    css: css/textbook.min.css
runtime: shiny_prerendered
description: >
  This interactive tutorial covers practice exercises for Week 3 of AHUM 1825. This includes filter(), select(), count(), and qplot().
---

<!-- Code for including shinyJS, currently used for toggling css classes toggleClass() | https://deanattali.com/shinyjs/advanced#usage-prerendered -->
```{r, echo=FALSE}
shiny::addResourcePath("shinyjs", system.file("srcjs", package = "shinyjs"))
```
```{r, context="server"}
shinyjs::useShinyjs(html = TRUE)
```
<script src="shinyjs/inject.js"></script>

<!-- End Code for shinyJs -->
```{r setup, include=FALSE}
library(learnr)
library(palmerpenguins)
library(ahum1825)
gradethis::gradethis_setup()

##Setup Homework Grade Tracker - keeps track of a student's name, grade, and answers to each homework problem
## WARNING: The default values are set inside the "session_start" event handler, which is inside of homework.Rmd
hmw_info <- NULL

#Uncomment to view all the Learnr events and their data. The data field is not well documented, so this a handy way to view everything
#new_recorder <- function(tutorial_id, tutorial_version, user_id, event, data) {
   #cat("Original: \n", tutorial_id, " (", tutorial_version, "): ", user_id, ", ", event, "\nData:\n", sep = "")
   #print(data)
#}
#options(tutorial.event_recorder = new_recorder)

library(tidyverse)
library(shinyjs)
library(reactable)

addResourcePath("SharedResources", "../SharedResources")#Add access to a shared set of images. Like the banner for the farm dataset

knitr::opts_chunk$set(error = FALSE, message = FALSE)
```

<!--{{dataset:FarmInventory}}-->

```{r context="server-start"}
##When the student begins a tutorial or hits "start over" this will reset their homework information. This will clear their
## homework grade and any homework problems submissions currently stored.
event_register_handler("session_start", function(session, event, data) {
  hmw_info <<- reactiveValues(s_id = get_tutorial_info()$user_id,
                           tutorial_name = "Week 3 Lab",
                           hmw_grade = 0.0, 
                           submissions = list(
                                  hmw1 = list(type="questionaire", 
                                               q_text="A multiple choice or true/false question",
                                               choices=c("Choice A", "Choice B", "Choice C"), 
                                               is_correct = FALSE, 
                                               submitted_answer = "",
                                               error_msg = ""),
                                  hmw2 = list(type="code", 
                                               q_text="This is a coding question with starter code ", 
                                               choices="starte_code_here",
                                               is_correct = FALSE, 
                                               submitted_answer = "",
                                               error_msg = ""),
                                  hmw3 = list(type="code", 
                                               q_text="This is a coding question without starter code ", 
                                               choices=NULL,
                                               is_correct = FALSE, 
                                               submitted_answer = "",
                                               error_msg = "")
                                   ))
})
```

<!-- 
Todo: 
  Create practice problems for each section
  Add banner images
  Update hmw_info to track each of the homework problems
  
-->

## Week 3 Lab

### <span style="opacity:0;">week_3_lab()</span>

![Cartoon showing a round monster holding up a sign that says Week 3 Lab”](images/Blank Banner.png){width=100%} 
<p align="center">"Artwork created with sample graphics provided by <a href="https://twitter.com/allison_horst?lang=en">@allison_horst</a>" </p>
  
For this week's recitation, we will be working through exercises that cover:

<dl>
  <dt>`filter()`</dt>
  <dd>This tidyverse function helps to select a subset of observations (rows) from a datset. This is handy both for exploring your dataset (e.g. how many observations (if any) fit your condition?) or feeding your dataset into a visualization (e.g. create a line graph, but only use the observations that meet your condition)</dd>
  <dt>`select()`</dt>
  <dd>This tidyverse function also provides a subset, but of variables (columns). This can make it easier to view your results, by only including the variables you are interested in.</dd>
  <dt>`count()`</dt>
  <dd>This tidyverse function counts the number of unique values that exist in a variable or a set of variables. This is a great way to see how much variety is in your dataset's variables</dd>
  <dt>`qplot()`</dt>
  <dd>qplot stands for *quick* plot and as its name suggests allows you to quickly create a plot from a given dataset. This will serve as our entry-point into creating visualizations in R, later on we'll look into how we can create more complex visualizations with `ggplot()`</dd>
</dl>

## Dataset  
 
```{r load_practice_set, child="../SharedResources/datasets/PalmerPenguins/description.Rmd" }

```  


<!-- ### <span style="opacity:0;">Quick Tip </span> -->

<!-- <div class="quick_note_box"> -->
<!--   <p class="quick_note_header">`r emo::ji("bookmark_tabs")` Quick Note</p> -->
<!--   <p> A helpful message for the reader:</p> -->
<!-- ```{r quick_note, R.options=list(width=25)} -->
<!-- #code example -->
<!-- simple -->
<!-- ``` -->
<!-- </div> -->

<!-- Tips to list: Edge cases, check for NA, incremental coding, debugging -->

## Filter  

![Cartoon showing three fuzzy monsters either selecting or crossing out rows of a data table. If the type of animal in the table is “otter” and the site is “bay”, a monster is drawing a purple rectangle around the row. If those conditions are not met, another monster is putting a line through the column indicating it will be excluded. Stylized text reads “dplyr::filter() - keep rows that satisfy your conditions.”](images/filter.jpg){width=100%} 
<p align="center">"Artwork by <a href="https://twitter.com/allison_horst?lang=en">@allison_horst</a>" </p>


#### Problem 1 

Fill in the blanks so that only the penguins from the Biscoe island are listed.
```{r week_3_hmw1, exercise = TRUE, exercise.blanks = "___+", }
filter(penguins, _______)
```

```{r week_3_hmw1-check}
grade_result(
  pass_if(~identical(.result, filter(penguins, island == "Biscoe")))#Replace TRUE with the solution's code
)
```

#### Problem 2 

Reverse your answer to problem 1, only list the penguins that are *not* on Biscoe island.
```{r week_3_hmw2, exercise = TRUE, exercise.blanks = "___+", }

```

```{r week_3_hmw2-check}
grade_result(
  pass_if(~identical(.result, filter(penguins, island != "Biscoe")))#Replace TRUE with the solution's code
)
```

#### Problem 3 

Debugging Exercise: Fix the code chunk so that only the small penguins, who weigh less than 3 kg, from 2007 and 2008 are listed.
```{r week_3_hmw3, exercise = TRUE}
filter(penguins, bodyMassG < 3000, year = 2007 | 2008)
```

```{r week_3_hmw3-check}
grade_result(
  pass_if(~identical(.result, filter(penguins, body_mass_g < 3000, (year == 2007 | year == 2008))))#Replace TRUE with the solution's code
)
```

## Select

#### Problem 4 

Fill in the blanks so that only the penguin's island, species, and flipper length are listed.
```{r week_3_hmw4, exercise = TRUE, exercise.blanks = "___+", }
select(penguins, ____, _____, ____)
```

```{r week_3_hmw4-check}
grade_result(
  pass_if(~identical(.result, select(penguins, island, species, bill_length_mm)))#Replace TRUE with the solution's code
)
```

#### Problem 5 

Write a select statement so that only the 4 columns are listed.
```{r week_3_hmw5, exercise = TRUE, exercise.blanks = "___+", }

```

```{r week_3_hmw5-check}
grade_result(
  pass_if(~identical(.result, select(penguins, 5:8)))#Replace TRUE with the solution's code
)
```

## Count

#### Problem 6 

```{r week_3_hmw6, exercise = TRUE, exercise.blanks = "___+", }

```

```{r week_3_hmw6-check}
grade_result(
  pass_if(~identical(.result, ))
)
```

#### Problem 7 

```{r week_3_hmw7, exercise = TRUE, exercise.blanks = "___+", }

```

```{r week_3_hmw7-check}
grade_result(
  pass_if(~identical(.result, ))
)
```

#### Problem 8 

```{r week_3_hmw8, exercise = TRUE, exercise.blanks = "___+", }

```

```{r week_3_hmw8-check}
grade_result(
  pass_if(~identical(.result, ))
)
```

## Quick Plot

#### Problem 9 

```{r week_3_hmw9, exercise = TRUE, exercise.blanks = "___+", exercise.reveal_solution = FALSE }

```

```{r week_3_hmw9-solution}
qplot(data=penguins, bill_length_mm, body_mass_g)
```

```{r week_3_hmw9-code-check}
grade_code()
```


#### Problem 10 

```{r week_3_hmw10, exercise = TRUE, exercise.blanks = "___+", exercise.reveal_solution = FALSE}

```

```{r week_3_hmw10-solution}

qplot(data=penguins, bill_length_mm, body_mass_g, color=islands)
```

```{r week_3_hmw10-code-check}
grade_code()
```

## Building Connections

#### Problem 11 

```{r week_3_hmw11, exercise = TRUE, exercise.blanks = "___+", }

```

```{r week_3_hmw11-check}
grade_result(
  pass_if(~identical(.result, ))
)
```

#### Problem 12 

```{r week_3_hmw12, exercise = TRUE, exercise.blanks = "___+", }

```

```{r week_3_hmw12-check}
grade_result(
  pass_if(~identical(.result, ))
)
```

#### Problem 13 

```{r week_3_hmw13, exercise = TRUE, exercise.blanks = "___+", }

```

```{r week_3_hmw13-check}
grade_result(
  pass_if(~identical(.result, ))
)
```

---

## Download Results

```{r, child="../SharedResources/grading/hmw_grader.Rmd"}
```

```{r, child="../SharedResources/grading/download_homework.Rmd"}
```
