---
title: "Using Select"
output:
  learnr::tutorial:
    progressive: true
    allow_skip: true
    fig_caption: false
    css: css/textbook.min.css
runtime: shiny_prerendered
description: >
  Learn how to pick a subset of columns in a tibble with `select()`.
---

<!-- Code for including shinyJS, currently used for toggling css classes toggleClass() | https://deanattali.com/shinyjs/advanced#usage-prerendered -->
```{r, echo=FALSE}
shiny::addResourcePath("shinyjs", system.file("srcjs", package = "shinyjs"))
```
```{r, context="server"}
shinyjs::useShinyjs(html = TRUE)
```
<script src="shinyjs/inject.js"></script>

<!-- End Code for shinyJs -->
```{r setup, include=FALSE}
library(tidyverse)
library(shinyjs)
library(reactable)
library(emo) # Emojis library
library(learnr)
gradethis::gradethis_setup()

##Setup Homework Grade Tracker
## WARNING: The default values are set inside the "session_start" event handler, which is inside of homework.Rmd
hmw_info <- NULL

#Uncomment to view all the Learnr events and their data
#new_recorder <- function(tutorial_id, tutorial_version, user_id, event, data) {
   #cat("Original: \n", tutorial_id, " (", tutorial_version, "): ", user_id, ", ", event, "\n", sep = "")
#}
#options(tutorial.event_recorder = new_recorder)

shiny::addResourcePath("SharedResources", "../SharedResources")#Needed for accessing images/other files

knitr::opts_chunk$set(error = FALSE, message = FALSE)
```

## What is select()?

### <span style="opacity:0;">select()</span>

![Cartoon showing a round monster holding up a sign that says dplyr::select”](images/tutorial_banner.png){width=100%} 
<p align="center">"Artwork created with sample graphics provided by <a href="https://twitter.com/allison_horst?lang=en">@allison_horst</a>" </p>
  
If you review the R Documentation for the [select() function](https://dplyr.tidyverse.org/reference/select.html) you'll find the following explanation:  

 > Select (and optionally rename) variables in a data frame, using a concise mini-language that makes it easy to refer to variables based on their name.  
 
In practice, we will use `select()` to help make large datasets more manageable, focusing only on the variables, or columns, we want to work with. `select()` can also be helpful when displaying our datasets, so that we can highlight key variables and improve readability by renaming the often shortened keywords used in datasets (e.g. avg_mpg -> Average Miles Per Gallon). Keep in mind, this is all done **without modifying or altering the original dataset!** 

## Practice Dataset  

```{r list-practice-dataset, child="../SharedResources/datasets/Mpg/description.Rmd"}
```

```{r, echo=FALSE}
head(mpg)
```


## Selecting Columns  

--- 

<span style="font-weight:bold; font-size:2em"> <span class="color1">subset</span> <- select(<span class="color2">data</span>, <span class="color3" >chosen_variables</span>)</span>

<dl>
  <dt class="color1 medium">subset</dt>
  <dd> `select()` doesn't actually modify the provided dataset, so you'll need to either store the results in order to use them in other operations or use the pipe **%>%** operator to pass the results onto another function.</dd>
  <dt class="color2 medium">data</dt>
  <dd>Here you provide the name of the dataset you want to work with. For most of our practice exercises, that will be the **plastic** dataset.</dd>
  <dt class="color3 medium">chosen_variables</dt>
  <dd>This is the most important part for using `select()`, the list of variables you want to keep or remove. This can be as simple as explicitly listing each variable you want to use or be a little more complex by using one of `select()`'s operators. The operators are defined below, along with code examples used on a simple dataset. 
    </dd>
</dl>  

--- 

```{r}
simple <- tibble(columnA = 24, columnB = 80, columnC = -3, columnD = 11) 
simple
```

--- 

###  {.tabset}

#### <span class="color3 medium bolder">name </span>
<p style="margin-top:15px"> Explicitly name the variable(s) you want to keep and their order.</p>
<table class="table">
  <tr>
    <td style="width:40%;"> `select(simple, columnC, columnA)`</td> 
    <td style="width:60%;">
      `r select(simple, columnC, columnA)`
    </td>
  </tr>
</table>

#### <span class="color3 medium bolder">position</span> 
<p style="margin-top:15px">List the position of the variable(s) within the dataset that you want to keep and their new order (i.e. 1 for first, 2 for second, 5 for fifth) </p>
<table class="table">
  <tr>
    <td style="width:40%;"> `select(simple, 2, 4, 1)`</td> 
    <td style="width:60%;">
      `r select(simple, 2, 4, 1)`
    </td>
  </tr>
</table>

#### <span class="color3 medium bolder">: (colon) </span> 
<p style="margin-top:15px">Select a range of variables to keep, using either their <span class="color3 bold">names</span> or <span class="color3 bold">positions</span> </p>
<table class="table">
  <tr>
    <td style="width:40%;"> `select(simple, columnB:columnD)`</td> 
    <td style="width:60%;">
      `r select(simple, columnB:columnD)`
    </td>
  </tr>
  <tr>
    <td> `select(simple, 1:3)`</td> 
    <td>
      `r select(simple, 1:3)`
    </td>
  </tr>
</table>

#### <span class="color3 medium bolder">\- (minus)</span> 
<p style="margin-top:15px">Choose a column to remove, using either its <span class="color3 bold">name</span> or <span class="color3 bold">position</span></p>
<table class="table">
  <tr>
    <td style="width:40%;"> `select(simple, -columnA)`</td> 
    <td style="width:60%;">
      `r select(simple, -columnA)`
    </td>
  </tr>
  <tr>
    <td> `select(simple, -3, -4)`</td> 
    <td>
      `r select(simple, -3, -4)`
    </td>
  </tr>
</table> 
### {-}
<div class="quick_note_box">
  <p class="quick_note_header">`r emo::ji("bookmark_tabs")` Quick Note</p>
  <p> If you ever forget the names of the variables in your dataset, you can use `colnames()` to quickly list them. For example, if you wanted to know the names of the variables in the mpg dataset you would write:</p>
```{r use_colnames, R.options=list(width=25)}
colnames(mpg)
```
</div>

<!-- Add widget here --> 

### <span style="opacity:0;">Blank</span>
### Practice Problems

#### Practice Problem 1
Finish the code block, so that only the variables for the car's model, year, and vehicle type are selected.
```{r select_practice1, exercise = TRUE, exercise.blanks = "___+"}
select(mpg, _____, _____, ____)
```

```{r select_practice1-solution}
select(mpg, model, year, class)
```

```{r select_practice1-check}
grade_result(
  pass_if(~identical(.result, select(mpg, model, year, class)))
)
```

#### Practice Problem 2
Write the necessary code so that only the 4th to the 8th variables are selected.
```{r select_practice2, exercise = TRUE}

```

```{r select_practice2-solution}
select(mpg, 4:8)
```

```{r select_practice2-check}
grade_result(
  pass_if(~identical(.result, select(mpg, 4:8)))
)
```

#### Practice Problem 3
Write the necessary code to remove the drive train and transmission type variables.
```{r select_practice3, exercise = TRUE}

```

```{r select_practice3-solution}
select(mpg, -drv, -trans)
```

```{r select_practice3-check}
grade_result(
  pass_if(~identical(.result, select(mpg, -drv, -trans)))
)
```

#### Practice Problem 4
Write the necessary code to reorder the variables so that the city and highway gas mileage and the fuel type are listed first. 
```{r select_practice4, exercise = TRUE}

```

```{r select_practice4-solution}
select(mpg, cty:fl, manufactorer:drv, class)
```


```{r select_practice4-check}
grade_result(
  pass_if(~identical(.result,select(mpg, cty:fl, manufactorer:drv, class)))
)
```

## Renaming Columns 

###  {.tabset}

--- 

<p style="font-weight:bold; font-size: 1.5em">select(data, <span class="color1" >new_variable_name</span> = <span class="color2">old_variable_name</span>)</p>

To rename a variable using `select()` all you have to do is provide a <span class="color1" >new variable name</span>, however the syntax can be a little awkward since the <span class="color1">new name</span> comes first, before the equals sign (=).

#### Rename with One Word   

<p style="margin-top: 15px">The easiest option for renaming a column is to use a single word. Make sure when you are renaming your variables that they still follow good coding practices (lowercase, noun phrase, meaningful).</p>

<table class="table">
  <tr>
    <th style="width: 50%;">Meaning*ful* Variable Names</th>
    <th style="width: 50%;">Meaning*less* Variable Names</th>
  </tr>
  <tr>
    <td>
```{r}
select(mpg, transmission = trans) %>%
  head(n=2)
``` 

  </td>
  <td>
```{r}
select(mpg, t = trans) %>%
  head(n=2)
``` 
  </td>
  </tr>
</table>
#### Rename with Multiple Words  
<p style="margin-top: 15px">Be careful when working with multiple words. At first you might try to write out a few words, but be careful variable names cannot contain spaces! Otherwise you'll get an error message like in the code chunk below. Instead make sure to separate your lowercase words with under_scores.</p>
```{r, error=TRUE}
select(mpg, drive train = drv)
```

<table class="table">
  <tr>
    <th style="width: 40%;">Meaning*ful* Variable Names</th>
    <th style="width: 60%;">Meaning*less* Variable Names</th>
  </tr>
  <tr>
    <td></td>
    <td>
      Long winded variable names are unwieldy and prone to typos. Also never use the actual values for a variable as its name! If new values are added or removed, the name would no longer accurately describe the variable's contents/purpose.
    </td>
  <tr>
    <td>
```{r}
select(mpg, drive_train = drv) %>%
  head(n=2)
```
  </td>
  <td>
```{r}
select(mpg, front_rear_or_four_wheel_drive = trans) %>%
  head(n=2)
``` 
  </td>
  </tr>
</table>

#### Be Careful of " " and .   
<p style="margin-top: 15px">When viewing code snippets online, you may see other styles for naming variables with multiple words. But be careful, while some are just different (e.g. CamelCasing) others can be downright problematic. Specifically, don't use dot.notation or "quotation marks".</p>
<table class="table">
  <tr>
    <th style="width: 50%;">Quotes **"** or **'**</th>
    <th style="width: 50%;">Dot **.** Notation</th>
  </tr>
  <tr>
    <td>
      Double quotes may appear more readable, but it will force all future operations to use quotes as well. For example, if you pipe `%>%` the results from select to filter, you'll need to reference your renamed columns using quotes "  ". In general, you should rename columns so that the quotes are removed.        
    </td>
    <td>
      Using dots, i.e. periods, to separate the words in your variable names used to be common practice in R and you can still find plenty of datasets that still follow Dot.Notation, like the `iris` dataset. This practice has fallen out of vogue because it resembles the same notation used by class methods in R. Classes are a topic beyond the scope of this tutorial, so to simplify just **don't use dot.notation**.
    </td>
  </tr>
  <tr>
    <td>
```{r}
select(mpg, "city mpg" = cty) %>%
  filter("city mpg" > 24) %>%
  head(n=2)
```
  </td>
  <td>
```{r}
select(mpg, Fuel.Type = fl) %>%
  head(n=2) 

``` 
  </td>
  </tr>
</table>

### Practice Problems
For each practice problem, ensure you are renaming each variable so that it follows our preferred naming conventions.

#### Practice Problem 5
Finish the code block so that the hwy variable is renamed as highway.
```{r select_practice5, exercise = TRUE, exercise.blanks = "___+"}
select(mpg, _____= hwy)
```

```{r select_practice5-solution}
select(mpg, highway = hwy)
```

```{r select_practice5-check}
grade_result(
  pass_if(~identical(.result, select(mpg, highway = hwy)))
)
```

#### Practice Problem 6
Finish the code block so that the class variable is renamed as car type.
```{r select_practice6, exercise = TRUE, exercise.blanks = "___+"}
select(mpg, _____= class)
```

```{r select_practice6-solution}
select(mpg, car_type = class)
```

```{r select_practice6-check}
grade_result(
  pass_if(~identical(.result, select(mpg, car_type = class)))
)
```

#### Practice Problem 7
Write the necessary code to list the variables manufacturer, model, and year. They should be followed by the renamed variables fuel type, transmission type, and highway.
```{r select_practice7, exercise = TRUE}

```

```{r select_practice7-solution}
select(mpg, manufacturer, model, year, fuel_type = fl, transmission_type = trans, highway = hwy)
```

```{r select_practice7-check}
grade_result(
  pass_if(~identical(.result, select(mpg, manufacturer, model, year, fuel_type = fl, transmission_type = trans, highway = hwy)))
)
```
## Homework Problems

<!-- 
  The description for the homework dataset
-->  

```{r include_hmw, child="data/hmw_data_description.Rmd"}
```

<!-- 
  Homework problems are stored in a separate file "homework_#.Rmd", where the # represents a specific set of homework problems.
  In order to prevent 
-->
```{r include_hmw, child="homework_1.Rmd"}
```
<!-- 
  The hmw_grader.Rmd child contains event listeners that evaluate a student's work each time they submit/answer a problem
-->
```{r, child="../SharedResources/grading/hmw_grader.Rmd"}
```
<!-- 
  The download_homework.Rmd child adds a button for students to download their homework submission.
  This report only includes the homework problems and none of the practice problems.
-->
```{r, child="../SharedResources/grading/download_homework.Rmd"}
```


## Helper Functions  

###  {.tabset}

--- 

Explicitly listing the variables you need can work for creating a quick subset but manually writing out all those variable names can get tedious real fast. If you need to make a complex selection of variables, turn to the helper functions!

#### contains()
<p style="margin-top: 15px">The `contains()` function will check each variable name to see if *contains* a literal string. 
This can be handy to grab a subset of variables that share a common substring.</p>  
**You can list a single substring** 
```{r}
select(green_vehicles, contains("mpg")) %>% #Retrieve the miles per gallon variables
  head(2) #limit to only 2 rows
```

**or multiple substrings.**
```{r}
select(green_vehicles, contains(c("mpg", "stnd"))) %>% #Retrieve the miles per gallon variables
  head(2) #limit to only 2 rows
```

#### starts_with()
<p style="margin-top: 15px">The `starts_with()` function works similar to `contains()`, but is a bit more focused. Instead of searching the entire variable name for a given substring, it only checks at the start of the name. In other words, its checking for a given prefix. </p>
**You can list a single prefix** 
```{r}
select(green_vehicles, starts_with("stnd")) %>% #Retrieve the miles per gallon variables
  head(2) #limit to only 2 rows
```

**or multiple prefixes.**
```{r}
select(green_vehicles, starts_with(c("stnd", "c"))) %>% #Retrieve the miles per gallon variables
  head(2) #limit to only 2 rows
```

#### ends_with()
<p style="margin-top: 15px">The `end_with()` function is the polar opposite of `starts_with()`. Instead of searching the for a given prefix, it checks the end of the name for given suffix. </p>  
**You can list a single suffix**
```{r}
select(green_vehicles, ends_with("score")) %>% #Retrieve the miles per gallon variables
  head(2) #limit to only 2 rows
```

**or multiple suffixes**
```{r}
select(green_vehicles, ends_with(c("ion", "_mpg"))) %>% #Retrieve the miles per gallon variables
  head(2) #limit to only 2 rows
```

#### num_range()
<p style="margin-top: 15px; font-size:1.5em;">num_range(<span class="color1" >"prefix"</span>, <span class="color2" >start</span>:<span class="color3" >end</span>)</p>
The `num_range()` function expands the functionality of `starts_with()`. It appends a given prefix with a series of numbers. This is a great way to search for columns in a generated dataset!
The [billboard dataset](https://tidyr.tidyverse.org/reference/billboard.html), packaged with tidyverse, lists the top 100 songs of year 2000 and provides each song's rank for each week after it entered the top 100. Working with `num_range()` we can easily grab a specific subset of weekly rankings.
```{r}
select(billboard, num_range("wk", 5:15)) %>% #Retrieve the rankings for week 5 to 15 
  head(2) #limit to only 2 rows
```


#### everything()
<p style="margin-top: 15px">As has been repeatedly stated, writing out the names of all the remaining variables in a data frame can be so tedious! Luckily, there the `everything()` function. It will grab all of the columns not explicitly named and place them at the end `select()`'s new data frame.</p>

```{r}
select(green_vehicles, model, year, smartway, everything()) %>% #Place all the remaining variables after smartway
  head(2) #limit to only 2 rows
```

#### dplyr::rename()  

<p style="margin-top: 15px">While select can be used for renaming columns, it does come with a big drawback : you need to list all of your columns to prevent them from being dropped or reordered.
If you just need to rename a column, without creating a subset, use dplyr's `rename()` function. Luckily, it has the same syntax as `select()` for changing names: </p>

<p style="font-weight:bold; font-size: 1.5em; text-align: center;margin-bottom:0px;">rename(data, <span class="color1" >new_variable_name</span> = <span class="color2">old_variable_name</span>)</p> 

<img  src="images/rename_sm.jpg" alt="A mouse changing data frames variable (column) names to 'nemesis' and 'lair' from 'species' and 'site', with news of a cheese heist in the background." style="width:80%; margin:0 10%" />

<p align="center">"Artwork by <a href="https://twitter.com/allison_horst?lang=en">@allison_horst</a>" </p>

```{r}
rename(mpg, drive_train = drv) %>% #rename drv as drive_train
  head(n=2) #limit the output to two rows
```


